---
title: 'SVD (Singular value decomposition)'
description: 'In this chapter, we introduce a method of decomposing matrices into  three components.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: 2272ff163d
xp: 50
```

`@projector_key`
0d1c02e84535d03a61017a969fbd7e27
